# W8 PIT Practical 

# Fibonacci number
---

The Fibonacci numbers are `0, 1, 1, 2, 3, 5, 8, 13, 21...`. 

The formula is `Fn = Fn-1 + Fn-2`, with `F0=0` and `F1=1`. 

For example,
- `F2 = F1 + F0 = 1 + 0 = 1`
-  `F3 = F2 + F1 = 1 + 1 = 2`
-  `F4 = F3 + F2 = 2 + 1 = 3`
- `F5 = F4 + F3 = 3 + 2 = 5`
- ...

## Task 1 (25%)
Please complete the function `fibo` below that displays Fibonacci number using a **non-recursive** approach

```python
def fibo(n):
    i = 1
    f0 = 0 #0th result
    f1 = 1 #1st result
    fi = 1 #fib result
    while i < n+1:
        print(f'f{i} = {fi}') #i = counter for while loop
        if n > 0: #check to see if n is larger than 0
        # INPUT YOUR CODE BELOW
        # approx 3 lines
            fi = f0 + f1 #fibonacci seq, (Fn = F(n-1)+F(n-2))
            i += 1 #increase counter by 1
            f0 = f1 #assign new f0 to f1
            f1 = fi #assign new f1 to fi

        
# =======================
# EXAMPLE TESTING CASES
fibo(5)
# output:
# f1 = 1
# f2 = 1
# f3 = 2
# f4 = 3
# f5 = 5
```

Output:
```
f1 = 1
f2 = 1
f3 = 2
f4 = 3
f5 = 5
```

## Task 2 (25%)
Please display Fibonacci number using a recursive approach

```python
def fibo_rec(n):
    #0 or 1 in fib seq does not need to be added or subtracted
    if n == 0 or n == 1: 
        #return original value
        return n 
    else:
        #e.g. 3-1 = 2(4th element), 3-2 = 1(2 and 3rd element)
        return fibo_rec(n-1) + fibo_rec(n-2) 
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
n = 10
print(f'f{n} is: {fibo_rec(n)}') # f10 is: 55

n = 20
print(f'f{n} is: {fibo_rec(n)}') # f20 is: 6765
```

Output:
```
f10 is: 55
f20 is: 6765
```

## Task 3 (25%)
Please implement the quick sort algorithm that 
1. sort the numbers in **ascending** order and
1. uses the **first** element as the pivot.

```python
import random

def qsort(arr):
    # YOUR CODE BELOW
    # ~ 2 lines, the base case    
    if len(arr) < 2:
        return arr #there are only 2 elements in this case so it is already sorted
    else:
        # Read two lines below carefully, you will implement the partition function
        lesser, greater = partition(arr)
        return qsort(lesser) + [arr[0]] + qsort(greater) 
    
def partition(arr):
    lesser = []
    greater = []
    pivot = arr[0]
    # YOUR CODE BELOW
    # ~ 5 lines, create two lists: lesser and greater
    for n in arr[1:]: #for n in array's 1st element onwards
        if n <= pivot: #if n is equal to or smaller than the pivot, append it to lesser array
            lesser.append(n)
        else:
            greater.append(n) #else, append to greater array
    return lesser, greater
    #partition is called by qsort
    #if the pivot is always at the front then we can just skip 0th element
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
A = [12, 23, 43, 21]
arr = qsort(A)
print(arr) #[12, 21, 23, 43]
```

Output:
```
[12, 21, 23, 43]
```

## Exercise 4 (25%)
Please implement the quick sort algorithm that 
1. sort the numbers in **descending** order and
1. uses **random** element as the pivot.

For simplicity, you can assume all number in the array are unique.

```python
import random

def qsort_dec(arr):
    # YOUR CODE BELOW
    # ~5 LINES
    # Copy and modify the qsort() function above
        if len(arr) < 2: # if length of array is smaller than 2
            return arr #there are only 2 elements in this case so it is already sorted
        else:
        # Read two lines below carefully, you will implement the partition function
            lesser, greater, pivot = partition_dec(arr)
            return qsort_dec(lesser) + [pivot] + qsort_dec(greater)
    
def partition_dec(arr):
    lesser = []
    greater = []
    # YOUR CODE BELOW    
    # ~ 6 lines, copy and update from Task 3
    
    rand = random.randint(0, len(arr) - 1) #randomization within the range of the array length by using the random library import
    
    pivot = arr[rand] #pivot is randomized from above
    arr.pop(rand)   #removes the random pivot to prevent duplicates
    for n in arr: #for n in array's 1st element onwards

        
        
        if n >= pivot: #if n is equal to or smaller than the pivot
            #also checking to see if there are any duplicates in the list
            lesser.append(n) #append to lesser array
            # print(lesser)
        else:
            greater.append(n) #else, append to greater array
             # print(greater)
    
    return lesser, greater, pivot
    

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
A = [5, 9, 12, 16, 23]
arr = qsort_dec(A)
print(arr) # [23, 16, 12, 9, 5]
```

Output:
```
[23, 16, 12, 9, 5]
```

```python

```

Tests for Task 1 (points: 1)
```
>>> fibo(10)
f1 = 1
f2 = 1
f3 = 2
f4 = 3
f5 = 5
f6 = 8
f7 = 13
f8 = 21
f9 = 34
f10 = 55
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 2 (points: 1)
```
>>> fibo_rec(5) == 5
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> fibo_rec(10) == 55
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 3 (points: 1)
```
>>> a = qsort([5, 3, 2, -5])
>>> a[3] ==5 
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 4 (points: 1)
```
>>> qsort_dec([98, 23, 43, 21, 53, 29])
[98, 53, 43, 29, 23, 21]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
