# COMP SCI 2009 Programming for IT - W07 - Practical

<!-- BEGIN QUESTION -->

# Q1

In IPIT, you have already coded a number guessing game in one of your practical assignment. In the assignment, you are the one guessing the number, while the program tells you if your guess is larger or smaller than the answer. In this exercise, you will do the opposite, i.e., the user inputs a number (between 0 ~ 999) and a number guessing AI continuously guesses the number (using a binary search method) according to the user response.

Below is an example. *words between asterisk* are user inputs. 
```
Please input a number (0~999):  *523*
Is the number 500 ? :  *larger*
Is the number 725 ? :  *smaller*
Is the number 612 ? :   *smaller*
...
Is the number 523? :   *yes*
```

```python
# the range of the number
left = 0
right = 999

# user input the answer
answer = int(input(f'Please answer a number(0~999)'))

# AI guessing loop starts
cnt = 0
while True:
    # instead of random guess, AI should use binary search
    # REPLACE the line below
    guess = (left + right)//2 # double / to avoid floating error
    response = input(f'Is the number {guess}?') #changes the input to an integer 
    
    # INSERT YOUR CODE BELOW    
    # Update left or right according to the response
    # ~ 7 lines
    if response == 'smaller': #if response is smaller than guess user inputs smaller
        right = guess
    elif response == 'larger': #if response is larger than guess user inputs larger
        left = guess
    elif response == 'yes': #if response is correct user inputs yes
        break

    # terminate loop for autograder
    cnt += 1
    if cnt >= 15:
        break
```

Output:
```
Please answer a number(0~999)499
Is the number 499?yes
```

<!-- END QUESTION -->

# Q2

Below is the pseudo code of an **insertion sort** algorithm. Note that in this pseudocode the index starts from `1` instead of `0` and you will need to modify the code accordingly. 
```
INSERTION_SORT(A)
for j = 2 to A.length
    key = A[j]
    // Insert A[j] into the sorted sequence A[1..j-1]
    i = j - 1
    while i > 0 and A[i] > key
        A[i+1] = A[i]
        i = i - 1
    A[i+1] = key
```

Please do the following
1. Read the pseudo code and explain how **each line of the code** works. (<font color='red'>Before you start Googling about **insertion sort**, please spend 10 minutes reading the pseudocode and make sense of it!</font>)
2. Implement the pseudocode using Python in the cell below.

```python
def insert_sort(A):
    for j in range(0, len(A)):
        # ~6 line
        key = A[j] #key to sort the list of A
        i = j - 1
        
        while i >= 0 and A[i] > key: #before sorting, while i is greater or is equal to 0 and i in list A is greater than 'key'
            A[i + 1] = A[i] #sorts the integer from smallest to largest
            i -= 1 #'i' will always be 1 integer lower than 'j' (line 5)
        A[i + 1] = key #finishes the rearrangement
        
    return A

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE

A = insert_sort([81, 14, 3, 94, 35, 31, 28, 17, 13, 86])
print(A) # [3, 13, 14, 17, 28, 31, 35, 81, 86, 94]
A = insert_sort([77, 3, 71, 25, 91, 83, 89, 69, 53, 28])
print(A) # [3, 25, 28, 53, 69, 71, 77, 83, 89, 91]
```

Output:
```
[3, 13, 14, 17, 28, 31, 35, 81, 86, 94]
[3, 25, 28, 53, 69, 71, 77, 83, 89, 91]
```

# Q3: Max Consecutive Identical Characters

Given a string, return the maximum number of consecutive identical characters.

**Example 1**

- Input: `aabbccccddd`<br>
- Output: 4<br>
- Comment: The input has 2 consecutive `a`, 2 `b`, 4 `c`, 3 `d`, thus the maximum number of consecutive identical characters is 4.

**Example 2**

- Input: `abcffgfffffgg`<br>
- Output: 5<br>
- Comment: Although there are 7 `f`, the longest consecutive substring is `fffff` as there is a `g` between `ff` and `fffff`. 
**Example 3**

- Input: `apple`<br>
- Output: 2<br>

```python
def consec_char(input_str):
    max_cnt = -1   
    cnt = 0 #counter
    cur = input_str[0]
    for i in range(len(input_str)):
        c = input_str[i]
        # hint: check if c is equal to cur. If not, update cur and check max_cnt
        # <10 lines
        if c == cur: #if c is equal to cur
            
            cnt += 1 # + 1 to counter
            if cnt > max_cnt:
                max_cnt = cnt # max_cnt is updated to new counter
               
            
        else:
            cnt = 1 #if counter doesn't meet requirements of the greater than if statement, it is set to 0 again
            
        cur = c   
    return max_cnt

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
c = consec_char('aabbccccddd')
print(c) # 4
c = consec_char('abcffgfffffgg')
print(c) # 5
```

Output:
```
4
5
```

# Q4: Longest Continuous Decreasing Subsequence

Given an unsorted array of integers nums, return the length of the longest continuous decreasing subsequence (i.e. subarray). The subsequence must be **strictly** decreasing.

Formally, a continuous decreasing subsequence is defined by two indices l and r (l < r) such that it is [nums[l], nums[l + 1], ..., nums[r - 1], nums[r]] and for each l <= i < r, nums[i] > nums[i + 1].

**Example 1** 
- input: `[3,5,4,3,2,7,1]`
- output: 4
- The longest continuous decreasing subsequence is `5432`. Note that although `54321` is the longest decreasing subsequence, it is not a continuous one as `7` is betweeen `2` and `1`

**Example 2**
- input: `[7,8,9,7,7,6,5,3,6]`
- output: 4
- The subsequence is [7, 6, 5, 3]. The numbers must be strictly decreasing, thus the 4th and 5th elements are both `7` and do not satisfy the requirement. 

**Example 3**
- input: `[2,4,8,6,5,8]`
- output: 3

**Example 4**
- input: `[1, 2, 3]`
- output: 1
- Special case when there is no continuous decreasing subsequence. 

```python
def lcds(array):
    max_cnt = 0
    cnt = 0
    cur = array[0]
    # INSERT YOUR CODE BELOW
    for i in range(len(array)):
        c = array[i]
        if  c < cur: # if c is smaller to cur
            cnt += 1 # + 1 to counter
            
        else:
            cnt = 1 #if counter doesn't meet requirements of the greater than if statement, it is set to 0 again
        cur = c
        if cnt > max_cnt: # if counter is larger than max counter
            max_cnt = cnt # max_cnt is updated to new counter
 
    return max_cnt
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
c = lcds([3,5,4,3,2,7,1])
print(c) # 4
c = lcds([7,8,9,7,7,6,5,3,6])
print(c) # 4
```

Output:
```
4
4
```

```python

```

Tests for Q2 (points: 25)
```
>>> insert_sort([-5, 10, 3])
[-5, 3, 10]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> insert_sort([-5, 13, -5, 10])
[-5, -5, 10, 13]
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q3 (points: 25)
```
>>> consec_char('aaaddc')
3
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> consec_char('abcwwww')
4
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q4 (points: 25)
```
>>> lcds([5, 4, 3])
3
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> lcds([4, 5, 4, 4, 3, 2, 1])
4
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> lcds([4, 5, 6])
1
```
failure_message: Incorrect answer, try again!
success_message: Good job!
